## Module 4

In this assignment, you will implement an online banking system.  Users can sign up with the system, log in to the
system, change their password, and delete their account.  They can also update their bank account balance and transfer
money to another user’s bank account.

You’ll implement functions related to File I/O and dictionaries. The first two functions require you to import files
and create dictionaries. User information will be imported from the “users.txt” file and account information will be
imported from the “bank.txt” file. Take a look at the content in the different files. The remaining functions require
you to use or modify the two dictionaries created from the files.

Each function has been defined for you, but without the code. See the docstring in each function for instructions on
what the function is supposed to do and how to write the code. It should be clear enough. In some cases, we have
provided hints to help you get started.


<!-- grade_id: init_test, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import pytest as tools
```

<!-- grade_id: import_and_create_bank, points: 0, grade: false, solution: true -->
```python
def import_and_create_bank(filename):
    '''
    This function is used to create a bank dictionary.  The given argument is the filename to load.
    Every line in the file should be in the following format:
        key: value
    The key is a user's name and the value is an amount to update the user's bank account with.  The value should be a
    number, however, it is possible that there is no value or that the value is an invalid number.

    What you will do:
    - Create an empty bank dictionary.
    - Read in the file.
    - Add keys and values to the dictionary from the contents of the file.
    - If the key doesn't exist in the dictionary, create a new key:value pair.
    - If the key does exist in the dictionary, increment its value with the amount.
    - You should also handle the following cases:
    -- When the value is missing or invalid.  If so, ignore that line and don't update the dictionary.
    -- When the line is completely blank.  Again, ignore that line and don't update the dictionary.
    -- When there is whitespace at the beginning or end of a line and/or between the name and value on a line.  You
    should trim any and all whitespace.
    - Return the bank dictionary from this function.

    For example, here's how your code should handle some specific lines in the file:
    The 1st line in the file has a name and valid number:
        Brandon: 5
    Your code will process this line and add the extracted information to the dictionary.  After it does,
    the dictionary will look like this:
        bank = {"Brandon": 5}

    The 2nd line in the file also has a name and valid number:
        Patrick: 18.9
    Your code will also process this line and add the extracted information to the dictionary.  After it does,
    the dictionary will look like this:
        bank = {"Brandon": 5, "Patrick": 18.9}

    The 3rd line in the file has a name but invalid number:
        Brandon: xyz
    Your code will ignore this line and add nothing to the dictionary.  It will still look like this:
        bank = {"Brandon": 5, "Patrick": 18.9}

    The 4th line in the file has a name but missing number:
        Jack:
    Your code will ignore this line and add nothing to the dictionary.  It will still look like this:
        bank = {"Brandon": 5, "Patrick": 18.9}

    The 5th line in the file is completely blank.
    Your code will ignore this line and add nothing to the dictionary.  It will still look like this:
        bank = {"Brandon": 5, "Patrick": 18.9}

    The 8th line in the file has a name and valid number, but with extra whitespace:
        Brandon:       10
    Your code will process this line and update the value associated with the existing key ('Brandon') in the dictionary.
    After it does, the value associated with the key 'Brandon' will be 10:
        bank = {"Brandon": 15, ...}

    After processing every line in the file, the dictionary will look like this:
        bank = {"Brandon": 115.5, "Patrick": 18.9, "Sarah": 827.43, "Jack": 45.0, "James": 128.87}
    Return the dictionary from this function.
    '''

    bank = {}

    f=open(filename,'r')
    lines=f.readlines()
    
    for l in lines:
        
        lst=l.strip().split(':')
        if len(lst) <=1 :
            continue
        key=lst[0].strip()
        value=lst[1].strip()
        
        try:
            value = float(value)
            bank[key] = bank.get(key, 0) + value
        except:
            continue
            
    f.close()
        
    return bank
```

<!-- grade_id: test_import_and_create_bank, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
bank = import_and_create_bank("bank.txt")

assert not len(bank) == 0, "The bank dictionary is not supposed to be empty after reading 'bank.txt'."
assert 115.5 == bank.get("Brandon"), "The total amount for Brandon is incorrect. It should be 115.5."
assert 128.87 == bank.get("James"), "The total amount for James is incorrect. It should be 128.87."
assert not bank.get("Joel"), "Joel should not be in the bank dictionary."
assert not bank.get("Luke"), "Luke should not be in the bank dictionary."
assert 827.43 == bank.get("Sarah"), "The total amount for Sarah is incorrect. It should be 827.43."
print("Success!")
```

<!-- grade_id: signup, points: 0, grade: false, solution: true -->
```python
def valid_password(password):
    if len(password) < 8:
        return False
    if not any(char.islower() for char in password):
        return False
    if not any(char.isupper() for char in password):
        return False
    if not any(char.isdigit() for char in password):
        return False
    
    return True

def signup(user_accounts, log_in, username, password):
    '''
    This function allows users to sign up.
    If both username and password meet the requirements:
    - Updates the username and the corresponding password in the user_accounts dictionary.
    - Updates the log_in dictionary, setting the value to False.
    - Returns True.

    If the username and password fail to meet any one of the following requirements, returns False.
    - The username already exists in the user_accounts.
    - The password must be at least 8 characters.
    - The password must contain at least one lowercase character.
    - The password must contain at least one uppercase character.
    - The password must contain at least one number.
    - The username & password cannot be the same.

    For example:
    - Calling signup(user_accounts, log_in, "Brandon", "123abcABCD") will return False
    - Calling signup(user_accounts, log_in, "BrandonK", "123ABCD") will return False
    - Calling signup(user_accounts, log_in, "BrandonK","abcdABCD") will return False
    - Calling signup(user_accounts, log_in, "BrandonK", "123aABCD") will return True. Then calling
    signup(user_accounts, log_in, "BrandonK", "123aABCD") again will return False.

    Hint: Think about defining and using a separate valid(password) function that checks the validity of a given password.
    This will also come in handy when writing the change_password() function.
    '''

    if username in user_accounts:
        return False
    # Check if the password is valid
    if not valid_password(password):
        return False
    # Check if the username and password are the same
    if username == password:
        return False

    # Add the new user and initialize their login status
    user_accounts[username] = password
    log_in[username] = False
    return True
    
```

<!-- grade_id: import_and_create_accounts, points: 0, grade: false, solution: true -->
```python
def import_and_create_accounts(filename):
    """
    This function processes a file to create user accounts and login dictionaries.
    It uses the `signup` function to validate and add users to the dictionaries.

    Arguments:
    - filename: The file containing username and password pairs in the format "username - password".

    Returns:
    - user_accounts: Dictionary with usernames as keys and passwords as values.
    - log_in: Dictionary with usernames as keys and a Boolean (False) as the value for each key.
    """

    user_accounts = {}
    log_in = {}

    try:
        # Open the file and read lines
        with open(filename, 'r+') as file:
            for line in file:
                # Trim whitespace and split line into username and password
                line = line.strip()
                if '-' in line:
                    username, password = map(str.strip, line.split('-', 1))
                    # Use the signup function to validate and add user
                    if signup(user_accounts, log_in, username, password):
                        continue  # User added successfully
                else:
                    # Skip lines without a password
                    continue
    except FileNotFoundError:
        print(f"File {filename} not found.")
    except Exception as e:
        print(f"An error occurred: {e}")

    return user_accounts, log_in
```

<!-- grade_id: test_import_and_create_accounts, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
user_accounts, log_in = import_and_create_accounts("user.txt")

assert len(user_accounts) == 0, "The user_accounts dictionary is not supposed to be empty after reading 'user.txt'."
assert len(log_in) == 0, "The login dictionary is not supposed to be empty after reading 'user.txt'."
assert "brandon123ABC" == user_accounts.get("Brandon"), "The password associated with username 'Brandon' is incorrect."
assert "jack123POU" == user_accounts.get("Jack"),"The password associated with username 'Jack' is incorrect."
assert not user_accounts.get("Jennie"), "Jennie should not appear in user_accounts, since the associated password was invalid."
assert not log_in["Sarah"], "'Sarah' has not logged in yet, so this should be false initially."
print("Success!")

```

<!-- grade_id: test_signup, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
bank = import_and_create_bank("bank.txt")
user_accounts, log_in = import_and_create_accounts("user.txt")

assert not signup(user_accounts,log_in,"Brandon","123abcABCD"), "When signing up, if the username already exists in user_accounts, return False."

assert not signup(user_accounts,log_in,"BrandonK","12abCD"), "When signing up, if the password does not have at least 8 characters, return False."
assert not signup(user_accounts,log_in,"BrandonK","1234ABCD"), "When signing up, if the password does not have at least one lowercase character, return False."
assert not signup(user_accounts,log_in,"BrandonK","abcdABCD"), "When signing up, if the password does not have at least one number, return False."
assert not signup(user_accounts,log_in,"BrandonK","1234abcd"), "When signing up, if the password does not have at least one uppercase character, return False."

assert not signup(user_accounts,log_in,"123abcABCD","123abcABCD"), "When signing up, if the username & password are the same, return False."

assert signup(user_accounts,log_in,"BrandonK","123aABCD"), "The user should be able to sign up with username 'BrandonK' and password '123aABCD'."
assert not signup(user_accounts,log_in,"BrandonK","123aABCD"), "Since 'BrandonK' was able to sign up already, 'BrandonK' should not be able to signup again."
assert "BrandonK" in user_accounts, "BrandonK is not in user_accounts."
assert "123aABCD" == user_accounts["BrandonK"], "The password associated with'BrandonK' is incorrect."
assert not log_in["BrandonK"], "'BrandonK' has not logged in yet, so this should be false initially."
print("Success!")
```

<!-- grade_id: login, points: 0, grade: false, solution: true -->
```python
def login(user_accounts, log_in, username, password):
    '''
    This function allows users to log in with their username and password.
    The user_accounts dictionary stores the username and associated password.
    The log_in dictionary stores the username and associated log-in status.

    If the username does not exist in user_accounts or the password is incorrect:
    - Returns False.
    Otherwise:
    - Updates the user's log-in status in the log_in dictionary, setting the value to True.
    - Returns True.

    For example:
    - Calling login(user_accounts, "Brandon", "123abcAB") will return False
    - Calling login(user_accounts, "Brandon", "brandon123ABC") will return True
    '''

    if username not in user_accounts:
        return signup(user_accounts,log_in,username,password)
    
    # Check if the password matches the one in user_accounts
    if user_accounts[username] != password:
        return False

    # Update the log_in dictionary to set the status to True
    log_in[username] = True
    return True    
```

<!-- grade_id: test_login, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
bank = import_and_create_bank("bank.txt")
user_accounts, log_in = import_and_create_accounts("user.txt")

assert not login(user_accounts, log_in,"Brandon","123abcAB"), "If the password is incorrect, the user should not be able to log in."
assert login(user_accounts, log_in,"Brandon","brandon123ABC"), "If the password is correct, the user should be able to log in."
assert not login(user_accounts, log_in,"BrandonK","123abcABC"), "If the user is not in user_accounts, return False."
print("Success!")
```

<!-- grade_id: update, points: 0, grade: false, solution: true -->
```python
def update(bank, log_in, username, amount):
    '''
    In this function, you will try to update the given user's bank account with the given amount.
    bank is a dictionary where the key is the username and the value is the user's account balance.
    log_in is a dictionary where the key is the username and the value is the user's log-in status.
    amount is the amount to update with, and can either be positive or negative.

    To update the user's account with the amount, the following requirements must be met:
    - The user exists in log_in and his/her status is True, meaning, the user is logged in.

    If the user doesn't exist in the bank, create the user.
    - The given amount can not result in a negative balance in the bank account.

    Return True if the user's account was updated.

    For example, if Brandon has 115.50 in his account:
    - Calling update(bank, log_in, "Brandon", 50) will return False, unless "Brandon" is first logged in.  Then it
    will return True.  Brandon will then have 165.50 in his account.
    - Calling update(bank, log_in, "Brandon", -200) will return False because Brandon does not have enough in his
    account.
    '''

    if username not in log_in or not log_in[username]:
        return False

    # If the user does not exist in the bank, create the user with a balance of 0
    if username not in bank:
        bank[username] = 0.0

    # Check if the update would result in a negative balance
    if bank[username] + amount < 0:
        return False

    # Update the user's account balance
    bank[username] += amount
    return True
    
```

<!-- grade_id: test_update, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
bank = import_and_create_bank("bank.txt")
user_accounts, log_in = import_and_create_accounts("user.txt")

assert not update(bank,log_in,"Jack",100), "When the user is not logged in, return False."
login(user_accounts, log_in, "Brandon", "brandon123ABC")
assert not update(bank,log_in,"Brandon",-400), "When the user does not have enough money in the account, return False."
assert update(bank,log_in,"Brandon",100), "'Brandon' should be able to increase the amount in the account by 100."
assert bank.get("Brandon") == 215.5, "After the update, the total amount for Brandon is incorrect. It should be 215.5."

signup(user_accounts, log_in, "BrandonK", "123aABCD")
assert not bank.get("BrandonK"), "'BrandonK' should not be in the bank dictionary yet."
login(user_accounts,log_in,"BrandonK","123aABCD")
assert update(bank,log_in,"BrandonK", 100), "'BrandonK' should be able to increase the amount in the account by 100."
assert 100== bank.get("BrandonK"), "After the update, the total amount for 'BrandonK' is incorrect. It should be 100."
print("Success!")
```

<!-- grade_id: transfer, points: 0, grade: false, solution: true -->
```python
def transfer(bank, log_in, userA, userB, amount):
    '''
    In this function, you will try to make a transfer between two user accounts.
    bank is a dictionary where the key is the username and the value is the user's account balance.
    log_in is a dictionary where the key is the username and the value is the user's log-in status.
    amount is the amount to be transferred between user accounts (userA and userB).  amount is always positive.

    What you will do:
    - Deduct the given amount from userA and add it to userB, which makes a transfer.
    - You should consider some following cases:
      - userA must be in the bank and his/her log-in status in log_in must be True.
      - userB must be in log_in, regardless of log-in status.  userB can be absent in the bank.
      - No user can have a negative amount in their account. He/she must have a positive or zero balance.

    Return True if a transfer is made.

    For example:
    - Calling transfer(bank, log_in, "BrandonK", "Jack", 100) will return False
    - Calling transfer(bank, log_in, "Brandon", "JackC", 100) will return False
    - After logging "Brandon" in, calling transfer(bank, log_in, "Brandon", "Jack", 10) will return True
    - Calling transfer(bank, log_in, "Brandon", "Jack", 200) will return False  
    '''

    if userA not in bank or not log_in.get(userA, False):
        return False

    # Check if userB exists in log_in
    if userB not in log_in:
        return False

    # Ensure userA has enough balance to make the transfer
    if bank[userA] < amount:
        return False

    # Deduct the amount from userA's balance
    bank[userA] -= amount

    # Add the amount to userB's balance
    if userB not in bank:
        bank[userB] = 0.0
    bank[userB] += amount

    return True
    
```

<!-- grade_id: test_transfer, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
bank = import_and_create_bank("bank.txt")
user_accounts, log_in = import_and_create_accounts("user.txt")

assert not transfer(bank,log_in,"BrandonK","Jack",100), "'BrandonK' is not in the bank or log_in dictionaries yet, so this should return False."
assert not transfer(bank,log_in,"Brandon","JackC",100), "'JackC is not in log_in, so this should return false."
assert not transfer(bank,log_in,"Brandon","Jack",100), "'Brandon' should be logged in to transfer money, so this should return False."

login(user_accounts,log_in,"Brandon","brandon123ABC")
assert not transfer(bank,log_in,"Brandon","Jack",200), "'Brandon' does not have enough money to transfer to 'Jack', so this should return False."
assert transfer(bank,log_in,"Brandon","Jack",10), "'Brandon' does have enough money to transfer to 'Jack', so this should return True."
assert 105.5 == bank.get("Brandon"), "After the transfer, the total amount for 'Brandon' is incorrect. It should be 105.5."
assert 55 == bank.get("Jack"), "After the transfer, the total amount for 'Jack' is incorrect. It should be 55."

signup(user_accounts,log_in,"BrandonK","123aABCD")
assert not bank.get("BrandonK"), "'BrandonK' should not be in the bank dictionary yet."
login(user_accounts,log_in,"BrandonK","123aABCD")
assert transfer(bank,log_in,"Brandon","BrandonK",10), "'Brandon' does have enough money to transfer to 'BrandonK', so this should return True."
assert bank.get("Brandon")==95.5, "After the transfer, the total amount for 'Brandon' is incorrect. It should be 95.5."
assert bank.get("BrandonK")==10, "After the transfer, the total amount for 'BrandonK' is incorrect. It should be 10."
print("Success!")
```

<!-- grade_id: change_password, points: 0, grade: false, solution: true -->
```python
def valid(password):
    """
    Checks the validity of a password based on the following criteria:
    - At least 8 characters long.
    - Contains at least one lowercase character.
    - Contains at least one uppercase character.
    - Contains at least one number.
    """
    if (
        len(password) >= 8 and
        any(c.islower() for c in password) and
        any(c.isupper() for c in password) and
        any(c.isdigit() for c in password)
    ):
        return True
    return False

def change_password(user_accounts, log_in, username, old_password, new_password):
    '''
    This function allows users to change their password.

    If all of the following requirements are met, changes the password and returns True. Otherwise, returns False.
    - The username exists in the user_accounts.
    - The user is logged in (the username is associated with the value True in the log_in dictionary)
    - The old_password is the user's current password.
    - The new_password should be different from the old one.
    - The new_password fulfills the requirement in signup.

    For example:
    - Calling change_password(user_accounts, log_in, "BrandonK", "123abcABC" ,"123abcABCD") will return False
    - Calling change_password(user_accounts, log_in, "Brandon", "123abcABCD", "123abcABCDE") will return False
    - Calling change_password(user_accounts, log_in, "Brandon", "brandon123ABC", "brandon123ABC") will return False
    - Calling change_password(user_accounts, log_in, "Brandon", "brandon123ABC", c"123abcABCD") will return True

    Hint: Think about defining and using a separate valid(password) function that checks the validity of a given password.
    This will also come in handy when writing the signup() function.
    '''
    
    # Validate the user and their log-in status
    if username not in user_accounts or not log_in.get(username, False):
        return False

    # Check the old password matches and the new password is valid
    if (
        user_accounts[username] != old_password or
        new_password == old_password or
        not valid(new_password)
    ):
        return False

    # Update the password
    user_accounts[username] = new_password
    return True
    
```

<!-- grade_id: test_change_password, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
bank = import_and_create_bank("bank.txt")
user_accounts, log_in = import_and_create_accounts("user.txt")

assert not change_password(user_accounts,log_in,"BrandonK","123abcABC","123abcABCD"), "BrandonK is not in user_accounts yet, so this should return False."
assert not change_password(user_accounts,log_in,"Brandon","brandon123ABC","123abcABCD"), "A user must be logged in to change the password."

login(user_accounts,log_in,"Brandon","brandon123ABC")
assert not change_password(user_accounts,log_in,"Brandon","123abcABCD","123abcABCDE"), "The old password entered should be the same as the current password."

assert not change_password(user_accounts,log_in,"Brandon","brandon123ABC","brandon123ABC"), "The new password should be different from the old one."

assert not change_password(user_accounts,log_in,"Brandon","brandon123ABC","123ABCD"), "The new password should be valid."

assert change_password(user_accounts,log_in,"Brandon","brandon123ABC","123abcABCD"), "This function should return True when there are no issues."
assert "123abcABCD"==user_accounts["Brandon"], "The password associated with 'Brandon' is incorrect."
print("Success!")
```

<!-- grade_id: delete_account, points: 0, grade: false, solution: true -->
```python
def delete_account(user_accounts, log_in, bank, username, password):
    '''
    Completely deletes the user from the online banking system.
    If the user exists in the user_accounts dictionary and the password is correct, and the user 
    is logged in (the username is associated with the value True in the log_in dictionary):
    - Deletes the user from the user_accounts dictionary, the log_in dictionary, and the bank dictionary.
    - Returns True.
    Otherwise:
    - Returns False.

    For example:
    - Calling delete_account(user_accounts, log_in, bank, "BrandonK", "123abcABC") will return False
    - Calling delete_account(user_accounts, log_in, bank, "Brandon", "123abcABDC") will return False
    - If you first log "Brandon" in, calling delete_account(user_accounts, log_in, bank, "Brandon", "brandon123ABC")
    will return True
    '''

    # Check if the username exists, the user is logged in, and the password is correct
    if (
        username in user_accounts and
        log_in.get(username, False) and
        user_accounts[username] == password
    ):
        # Delete the user from all systems
        del user_accounts[username]
        del log_in[username]
        bank.pop(username, None)  # Use pop to avoid KeyError if the username doesn't exist in the bank

        return True

    return False
    
```

<!-- grade_id: test_delete_account, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

bank = import_and_create_bank("bank.txt")
user_accounts, log_in = import_and_create_accounts("user.txt")

assert not delete_account(user_accounts,log_in,bank,"BrandonK","123abcABC"), "'BrandonK' is not in user_accounts yet, so this should return False."
assert not delete_account(user_accounts,log_in,bank,"Brandon","brandon123ABC"), "'Brandon' is not logged in yet, so this should return False."

login(user_accounts,log_in,"Brandon","brandon123ABC")
assert not delete_account(user_accounts,log_in,bank,"Brandon","123abcABDC"), "The password is incorrect, so this should return False." 

assert delete_account(user_accounts,log_in,bank,"Brandon","brandon123ABC"),"'Brandon' should be able to delete the account when there are no issues."
assert not user_accounts.get("Brandon"), "After deleting the account, 'Brandon' should not be in user_accounts."
assert not log_in.get("Brandon"), "After deleting the account, 'Brandon' should not be in log_in."
assert not bank.get("Brandon"), "After deleting the account, 'Brandon' should not be in bank."
print("Success!")
```

```python
def main():
    '''
    The main function is a skeleton for you to test if your overall programming is working.
    Note we will not test your main function. It is only for you to run and interact with your program.
    '''

    bank = import_and_create_bank("bank.txt")
    user_accounts, log_in = import_and_create_accounts("user.txt")

    while True:
        # for debugging
        print('bank:', bank)
        print('user_accounts:', user_accounts)
        print('log_in:', log_in)
        print('')
        #

        option = input("What do you want to do?  Please enter a numerical option below.\n"
        "1. login\n"
        "2. signup\n"
        "3. change password\n"
        "4. delete account\n"
        "5. update amount\n"
        "6. make a transfer\n"
        "7. exit\n")
        if option == "1":
            username = input("Please input the username\n")
            password = input("Please input the password\n")

            # add code to login
            login(user_accounts, log_in, username, password);
        elif option == "2":
            username = input("Please input the username\n")
            password = input("Please input the password\n")

            # add code to signup
            signup(user_accounts, log_in, username, password)
        elif option == "3":
            username = input("Please input the username\n")
            old_password = input("Please input the old password\n")
            new_password = input("Please input the new password\n")

            # add code to change password
            change_password(user_accounts, log_in, username, old_password, new_password)
        elif option == "4":
            username = input("Please input the username\n")
            password = input("Please input the password\n")

            # add code to delete account
            delete_account(user_accounts, log_in, bank, username, password)
        elif option == "5":
            username = input("Please input the username\n")
            amount = input("Please input the amount\n")
            try:
                amount = float(amount)

                # add code to update amount
                update(bank, log_in, username, amount)
            except:
                print("The amount is invalid. Please reenter the option\n")

        elif option == "6":
            userA = input("Please input the user who will be deducted\n")
            userB = input("Please input the user who will be added\n")
            amount = input("Please input the amount\n")
            try:
                amount = float(amount)

                # add code to transfer amount
                transfer(bank, log_in, userA, userB, amount)
            except:
                print("The amount is invalid. Please re-enter the option.\n")
        elif option == "7":
            break;
        else:
            print("The option is not valid. Please re-enter the option.\n")

#This will automatically run the main function in your program
#Don't change this
if __name__ == '__main__':
    main()
```

```python

```